## Required Assignment 1.3: Selecting Data in Multiple Ways

**Expected Time: 45 Minutes**

**Total Points: 20 Points**

In this assignment, you will practice selecting data using `loc[]`, `iloc[]`, and `split`.

## Index:

- [Part 1](#Part-one)
- [Part 2](#Part-two)
- [Part 3](#Part-three)

```python
import pandas as pd
```

[Back to top](#Index:) 

## Part one

### Using `loc[]` to select data

**2 Points each**

In this part of the assignment, return the required information from the `data` DataFrame using the `loc[]` method.

```python
#create the DataFrame utilizing Pandas 
data = pd.read_csv('data/nba.csv')
data = data.iloc[:-1 , :]
data.head(5)
```

Use the `loc[]` method to retrieve the first row of the `data` DataFrame and assign the result to the `locinfo` variable.

<!-- grade_id: cell-b045abc8cbfc7df8, points: 0, grade: false, solution: true -->
```python
### GRADED


locinfo = None

### BEGIN SOLUTION
locinfo = data.loc[0]
### END SOLUTION

# ANSWER CHECK
print(locinfo)
```

Use the `loc[]` method on the `data` DataFrame to retrieve the information about LeBron James' `Name` and `Team`. Assign the result to the `teamdata` variable.

<!-- grade_id: cell-2ce96b09f45395ce, points: 0, grade: false, solution: true -->
```python
### GRADED


teamdata = None

### BEGIN SOLUTION
teamdata = data.loc[data.Name == 'LeBron James', 'Name':'Team']
### END SOLUTION

# ANSWER CHECK
print(teamdata)
```

[Back to top](#Index:) 

## Part two

### Using `iloc[]` to select data

In this part of the assignment, return the required information from the `data` DataFrame using the `iloc[]` method.

**2 Points each**

Use the `iloc[]` method to retrieve the first row of the `data` DataFrame and assign the result to the `fr` variable.

<!-- grade_id: cell-ac4f90f3ca044a80, points: 0, grade: false, solution: true -->
```python
### GRADED

fr = None

### BEGIN SOLUTION

fr = data.iloc[0]

### END SOLUTION

# ANSWER CHECK
print(fr)
```

Use the `iloc[]` method to retrieve the second row of the `data` DataFrame and assign the result to the `sr` variable.

<!-- grade_id: cell-c96cb0193a5accb4, points: 0, grade: false, solution: true -->
```python
### GRADED


sr = None

### BEGIN SOLUTION

sr =  data.iloc[1]

### END SOLUTION

# ANSWER CHECK
print(sr)
```

Use the `iloc[]` method to retrieve the last row of the `data` DataFrame and assign the result to the `lr` variable.

<!-- grade_id: cell-1843d8fe5e08d07a, points: 0, grade: false, solution: true -->
```python
### GRADED


lr = None

### BEGIN SOLUTION

lr =  data.iloc[-1] 

### END SOLUTION

# ANSWER CHECK
print(lr)
```

Use the `iloc[]` method to retrieve the first column of the `data` DataFrame and assign the result to the `fc` variable.

<!-- grade_id: cell-b2ddae2ebcad6705, points: 0, grade: false, solution: true -->
```python
###GRADED

 
fc = None

### BEGIN SOLUTION

fc =  data.iloc[:,0] 

### END SOLUTION

# ANSWER CHECK
print(fc)
```

Use the `iloc[]` method to retrieve the second column of the `data` DataFrame and assign the result to the `sc` variable.

<!-- grade_id: cell-35925d6c39ff9c32, points: 0, grade: false, solution: true -->
```python
### GRADED


sc = None

### BEGIN SOLUTION

sc =  data.iloc[:,1]

### END SOLUTION

# ANSWER CHECK
print(sc)
```

Use the `iloc[]` method to retrieve the last column of the `data` DataFrame and assign the result to the `lc` variable.

<!-- grade_id: cell-04a5a0d17f2b4901, points: 0, grade: false, solution: true -->
```python
### GRADED


lc = None

### BEGIN SOLUTION

lc =  data.iloc[:,-1] 

### END SOLUTION

# ANSWER CHECK
print(lc)
```

<!-- grade_id: cell-86adba4835e0586a, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
lc1 =  data.iloc[:,-1]
assert type(lc) == type(lc1)
lc == lc1

### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Part three

### Using `split()` to select data

In this part of the assignment, return the required information from the `data` DataFrame using the `split()` function.

**2 Points each**

In this exercise, use the `split()` function to return the first name of the first player in the `data` DataFrame. Assign the result to the variable `firstname`.



<!-- grade_id: cell-268b6113422e9d31, points: 0, grade: false, solution: true -->
```python
### GRADED



firstname = None

### BEGIN SOLUTION

firstname = data['Name'].str.split().str[0].iloc[0]

### END SOLUTION

# ANSWER CHECK
print('The first name is ' + firstname)
```

In this exercise, use the `split()` function to return the last name of the first player in the `data` DataFrame. Assign the result to the variable `lastname`.

<!-- grade_id: cell-66242b355c8d6e79, points: 0, grade: false, solution: true -->
```python
###GRADED

# Please print the last full name of the first series in the data frame.

lastname = None

### BEGIN SOLUTION
lastname =  data['Name'].str.split().str[1].iloc[0]
### END SOLUTION

# ANSWER CHECK
print('The last name is ' + lastname)
```

```python

```